Model loaded in 68 ms
i: 30 0 1 3 0 0 1 0 0 2 0
a: 30
i: 29 0 1 2 0 0 1 0 0 2 0
a: 29
i: 31 0 1 2 0 0 2 0 0 1 0
a: 31
i: 32 0 1 2 0 0 2 0 0 1 0
a: 32
i: 33 0 1 2 0 0 255 0 0 1 0
a: 33
i: 34 0 4 11 0 0 255 0 0 1 0
a: 34
i: 35 0 2 6 0 0 1 0 0 1 0
a: 35
i: 36 0 2 6 0 0 255 0 0 1 0
a: 36
i: 37 0 2 6 0 0 255 0 0 1 0
a: 37
i: 38 0 2 6 0 0 255 0 0 1 0
a: 38
i: 39 0 2 6 0 0 255 0 0 1 0
a: 39
i: 40 0 2 6 0 0 2 0 0 1 0
a: 40
i: 41 0 1 4 0 0 2 0 0 2 0
a: 41
i: 42 0 1 3 0 0 255 0 0 1 0
a: 42
i: 43 0 1 3 0 1 0 0 0 1 0
a: 43
s: 33
s: 30
s: 41
s: 42
s: 40
s: 37
s: 36
s: 34
s: 39
i: 45 1 2 6 0 0 255 0 0 2 0
a: 45
s: 38
i: 46 1 2 6 0 0 2 0 0 2 0
a: 46
i: 47 1 1 4 0 0 2 0 0 3 0
a: 47
s: 31
s: 46
i: 49 2 1 4 0 0 2 0 0 4 0
a: 49
s: 29
i: 58 1 1 3 1 0 255 0 0 3 0
a: 58
s: 35
i: 64 1 1 3 0 0 1 0 0 3 0
a: 64
s: 47
i: 65 2 1 3 0 0 2 0 0 5 0
a: 65
s: 43
s: 49
i: 66 3 1 3 0 0 2 0 0 6 0
a: 66
s: 65
s: 64
s: 58
i: 68 1 1 3 1 0 1 0 0 3 0
a: 68
r: 68
a: 30
a: 29
a: 64
i: 62 1 1 3 1 0 255 0 0 3 0
a: 62
i: 57 1 1 3 1 0 255 0 0 3 0
a: 57
s: 66
s: 30
i: 74 1 1 3 0 0 1 0 0 3 0
a: 74
s: 57
s: 62
i: 75 0 1 3 1 0 1 0 0 2 0
a: 75
s: 64
i: 77 1 1 3 1 0 1 0 0 3 0
a: 77
s: 29
i: 78 0 1 2 1 0 1 0 0 2 0
a: 78
s: 74
i: 79 1 1 3 1 0 1 0 0 3 0
a: 79
s: 79
s: 77
s: 75
s: 32
s: 45
r: 78
a: 74
a: 29
a: 64
a: 30
i: 61 1 1 3 1 0 255 0 0 3 0
a: 61
i: 54 1 1 3 1 0 255 0 0 3 0
a: 54
s: 29
s: 30
i: 88 2 1 3 0 0 1 0 0 4 0
a: 88
i: 89 1 1 3 0 0 1 0 0 3 0
a: 89
s: 61
s: 74
s: 88
s: 89
s: 64
s: 54
i: 90 2 1 3 1 0 1 0 0 4 0
a: 90
i: 91 1 1 3 1 0 1 0 0 3 0
a: 91
i: 92 0 1 3 1 0 1 0 0 2 0
a: 92
i: 94 2 1 3 2 0 255 0 0 6 0
a: 94
i: 95 2 1 3 2 0 255 0 0 6 0
a: 95
s: 91
i: 96 2 1 3 1 0 1 0 0 4 0
a: 96
s: 90
s: 94
s: 92
s: 95
s: 96
% Refutation found. Thanks to Tanya!
% SZS status Theorem for PUZ001+1
% SZS output start Proof for PUZ001+1
1. ? [X0] : (killed(X0,agatha) & lives(X0)) [input]
5. ! [X0] : (lives(X0) => (charles = X0 | butler = X0 | agatha = X0)) [input]
6. ! [X0,X1] : (killed(X0,X1) => hates(X0,X1)) [input]
7. ! [X0,X1] : (killed(X0,X1) => ~richer(X0,X1)) [input]
8. ! [X0] : (hates(agatha,X0) => ~hates(charles,X0)) [input]
9. ! [X0] : (butler != X0 => hates(agatha,X0)) [input]
10. ! [X0] : (~richer(X0,agatha) => hates(butler,X0)) [input]
11. ! [X0] : (hates(agatha,X0) => hates(butler,X0)) [input]
12. ! [X0] : ? [X1] : ~hates(X0,X1) [input]
13. agatha != butler [input]
14. killed(agatha,agatha) [input]
15. ~killed(agatha,agatha) [negated conjecture 14]
16. ~killed(agatha,agatha) [flattening 15]
17. ! [X0] : ((charles = X0 | butler = X0 | agatha = X0) | ~lives(X0)) [ennf transformation 5]
18. ! [X0] : (charles = X0 | butler = X0 | agatha = X0 | ~lives(X0)) [flattening 17]
19. ! [X0,X1] : (hates(X0,X1) | ~killed(X0,X1)) [ennf transformation 6]
20. ! [X0,X1] : (~richer(X0,X1) | ~killed(X0,X1)) [ennf transformation 7]
21. ! [X0] : (~hates(charles,X0) | ~hates(agatha,X0)) [ennf transformation 8]
22. ! [X0] : (hates(agatha,X0) | butler = X0) [ennf transformation 9]
23. ! [X0] : (hates(butler,X0) | richer(X0,agatha)) [ennf transformation 10]
24. ! [X0] : (hates(butler,X0) | ~hates(agatha,X0)) [ennf transformation 11]
25. ? [X0] : (killed(X0,agatha) & lives(X0)) => (killed(sK0,agatha) & lives(sK0)) [choice axiom]
26. killed(sK0,agatha) & lives(sK0) [skolemisation 1,25]
27. ! [X0] : (? [X1] : ~hates(X0,X1) => ~hates(X0,sK1(X0))) [choice axiom]
28. ! [X0] : ~hates(X0,sK1(X0)) [skolemisation 12,27]
29. lives(sK0) [cnf transformation 26]
30. killed(sK0,agatha) [cnf transformation 26]
34. ~lives(X0) | butler = X0 | agatha = X0 | charles = X0 [cnf transformation 18]
35. ~killed(X0,X1) | hates(X0,X1) [cnf transformation 19]
36. ~richer(X0,X1) | ~killed(X0,X1) [cnf transformation 20]
37. ~hates(charles,X0) | ~hates(agatha,X0) [cnf transformation 21]
38. hates(agatha,X0) | butler = X0 [cnf transformation 22]
39. richer(X0,agatha) | hates(butler,X0) [cnf transformation 23]
40. ~hates(agatha,X0) | hates(butler,X0) [cnf transformation 24]
41. ~hates(X0,sK1(X0)) [cnf transformation 28]
42. agatha != butler [cnf transformation 13]
43. ~killed(agatha,agatha) [cnf transformation 16]
45. ~killed(X0,agatha) | hates(butler,X0) [resolution 39,36]
46. hates(butler,X0) | butler = X0 [resolution 38,40]
49. butler = sK1(butler) [resolution 46,41]
52. 1 <=> charles = sK0 [avatar definition]
54. charles = sK0 <- (1) [avatar component clause 52]
56. 2 <=> agatha = sK0 [avatar definition]
57. agatha != sK0 <- (~2) [avatar component clause 56]
58. agatha = sK0 <- (2) [avatar component clause 56]
60. 3 <=> butler = sK0 [avatar definition]
62. butler = sK0 <- (3) [avatar component clause 60]
66. ~hates(butler,butler) [superposition 41,49]
70. killed(agatha,agatha) <- (2) [backward demodulation 30,58]
71. $false <- (2) [subsumption resolution 70,43]
72. ~2 [avatar contradiction clause 71]
74. hates(sK0,agatha) [resolution 30,35]
75. killed(butler,agatha) <- (3) [backward demodulation 30,62]
82. hates(butler,butler) <- (3) [resolution 45,75]
83. $false <- (3) [subsumption resolution 82,66]
84. ~3 [avatar contradiction clause 83]
85. butler = sK0 | agatha = sK0 | charles = sK0 [resolution 29,34]
86. butler = sK0 | charles = sK0 <- (~2) [subsumption resolution 85,57]
87. 1 | 3 | 2 [avatar split clause 86,56,60,52]
91. hates(charles,agatha) <- (1) [backward demodulation 74,54]
96. ~hates(agatha,agatha) <- (1) [resolution 91,37]
99. agatha = butler <- (1) [resolution 96,38]
100. $false <- (1) [subsumption resolution 99,42]
101. ~1 [avatar contradiction clause 100]
102. $false [avatar sat refutation 72,84,87,101]
% SZS output end Proof for PUZ001+1
% ------------------------------
% Version: Vampire 4.7 (commit ce51621ed on 2022-08-24 17:46:27 +0200)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.090 s
% ------------------------------
----  Runtime statistics ----
clauses created: 66
clauses deleted: 8
ssat_new_components: 3
ssat_reused_components: 2
ssat_sat_clauses: 5
total_frozen: 12
total_unfrozen: 7
unworthy child removed: 14
-----------------------------
% ------------------------------
